## ข้อสอบปฏิบัติ: การคงอยู่ของพนักงานบริษัท

บริษัทแห่งหนึ่งมีพนักงานลาออกเป็นจำนวนมากในช่วง 1 ปีที่ผ่านมา ผู้บริหารมีความกังวลและต้องการทราบถึงปัจจัยที่ส่งผลต่อการลาออกของพนักงาน จึงได้มอบหมายให้ท่านนำข้อมูลพนักงานมาวิเคราะห์และรายงานผล โดยฝ่ายบริหารงานบุคคลได้รวบรวมข้อมูลพนักงานจำนวน 1470 คนไว้ในไฟล์ `employee_quit.csv` ในไฟล์ดังกล่าวประกอบด้วยสถานะการลาออก ข้อมูลทั่วไป เช่น อายุ เพศ สาขาความเชี่ยวชาญ รวมถึงข้อมูลการงาน เช่น เงินเดือน ระยะเวลาที่ทำงานกับบริษัท ความพึงพอใจต่องาน การทำงานเกินเวลา เป็นต้น ท่านจะต้องนำไฟล์ดังกล่าวมาประมวลและวิเคราะห์ดังต่อไปนี้ด้วยภาษา Python

### Part 1: นำเข้าข้อมูลจากไฟล์ .csv (1 คะแนน) 

(PC2.1) จงนำเข้าข้อมูลจากไฟล์ employee_quit.csv โดยใช้คำสั่งจากไลบรารี Pandas 
<br>**สำคัญ:** ให้เก็บข้อมูลในตัวแปรชื่อ `df`

<!-- grade_id: read_csv, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

รันโค้ดต่อไปนี้เพื่อตรวจสอบความถูกต้องของตัวแปร `df` ของท่าน

<!-- grade_id: read_csv_test, points: 1, grade: true, solution: false -->
```python
"""ตรวจสอบชนิดของตัวแปร df และจำนวนคอลัมน์ของข้อมูลในเบื้องต้น"""
import pandas as pd
assert isinstance(df, pd.DataFrame)
assert len(df.columns) == 15
```

---

### Part 2: การสำรวจข้อมูล (7 คะแนน) 

(PC2.1) จงนำข้อมูลของพนักงาน 7 แรกใน `df` มาเก็บไว้ในตัวแปรชื่อ `df_top`  

<!-- grade_id: df_top, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

ตรวจสอบความถูกต้องของตัวแปร `df_top` ในเบื้องต้น

<!-- grade_id: df_top_test, points: 1, grade: true, solution: false -->
```python
"""ตรวจสอบว่า df_top มี 7 แถว"""
assert isinstance(df_top, pd.DataFrame)
assert len(df_top) == 7
```

---

(PC2.3) จากการตรวจสอบพบว่าคอลัมน์ `MonthlyIncome` ใน `df` มีข้อมูลสูญหาย (missing values) เป็นจำนวนมาก จง implement ฟังก์ชัน `count_missing_salary(df)` สำหรับนับจำนวนข้อมูลที่สูญหายในคอลัมน์ `MonthlyIncome` โดยฟังก์ชันจะต้อง return จำนวนข้อมูลที่สูญหาย 

**หมายเหตุ** ระบบจะตรวจสอบว่าฟังก์ชันของท่านมีการนับจริง หาก return เป็นค่าคงที่ ท่านจะไม่ได้คะแนน

<!-- grade_id: count_miss, points: 0, grade: false, solution: true -->
```python
def count_missing_salary(df):
    # YOUR CODE HERE
    raise NotImplementedError()
```

ตรวจสอบว่าฟังก์ชันสามารถนับจำนวนข้อมูลสูญหายได้ถูกต้อง

```python
count_missing_salary(df)
```

<!-- grade_id: count_miss_test, points: 2, grade: true, solution: false -->
```python
assert count_missing_salary(df) == 57 
```

---

(PC1.2, 2.5) เนื่องจากข้อมูลในคอลัมน์ `MonthlyIncome` มีค่าผิดปกติ (outliers) บางค่า จง implement ฟังก์ชัน `find_outliers_iqr(df)` สำหรับหาค่าผิดปกติด้วยวิธี Interquartile Range (IQR) โดยฟังก์ชันจะต้อง return เป็น list ของค่าเงินเดือนที่ผิดปกติทั้งหมด 

**หมายเหตุ** ระบบจะตรวจสอบว่าฟังก์ชันของท่านมีการหาค่าผิดปกติจริง หาก return เป็น list ของค่าคงที่ ท่านจะไม่ได้คะแนน

<!-- grade_id: outlier, points: 0, grade: false, solution: true -->
```python
def find_outlier_iqr(df):
    # YOUR CODE HERE
    raise NotImplementedError()
```

ตรวจสอบว่าฟังก์ชันสามารถหาค่าผิดปกติได้ถูกต้อง (แสดง 5 ค่าแรกที่พบ)

```python
find_outlier_iqr(df)[:5]
```

<!-- grade_id: outlier_test, points: 4, grade: true, solution: false -->
```python
assert len(find_outlier_iqr(df)) == 90 
```

---

### Part 3: การทำความสะอาดข้อมูล (7 คะแนน) 

(PC2.3, 2.5) จงเติมค่าที่สูญหายในคอลัมน์ `MonthlyIncome` ของ `df` ด้วยค่าเงินเดือนเฉลี่ยของพนักงาน แล้วบันทึกข้อมูลที่เติมค่าสูญหายแล้วไว้ในตัวแปรชื่อ `df_filled`

<!-- grade_id: impute, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

ตรวจสอบว่าข้อมูลเงินเดือนได้ถูกเติมด้วยค่าเฉลี่ยเรียบร้อย

<!-- grade_id: impute_test, points: 2, grade: true, solution: false -->
```python
assert df_filled['MonthlyIncome'].count() == 1470
```

---
(PC2.4) สาขาความเชี่ยวชาญของพนักงาน (`EducationField`) เป็นปัจจัยหนึ่งที่อาจใช้ทำนายการลาออกของพนักงานได้ แต่เนื่องจาก `EducationField` เป็นข้อมูลเชิงคุณลักษณะ ท่านจะต้องแปลงข้อมูลดังกล่าวให้เป็นเชิงปริมาณก่อนนำไปใช้ฝึกแบบจำลอง 

จงแปลงข้อมูล `EducationField` ใน `df` ด้วยวิธี one-hot encoding แล้วเก็บผลไว้ในตัวแปรชนิด DataFrame ชื่อ `df_field` และให้แต่ละคอลัมน์มีชื่อขึ้นต้นด้วยคำว่า `Field_` เช่น `Field_Marketing`, `Field_Medical` เป็นต้น

**หมายเหตุ** สำหรับข้อนี้ `df_field` จะเก็บเฉพาะข้อมูล `EducationField` ที่เป็นแบบ one-hot ท่านยังไม่ต้องนำข้อมูลนี้ไปแทนที่คอลัมน์เดิมใน `df`

<!-- grade_id: onehot, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

ตรวจสอบความถูกต้องของ `df_field` ในเบื้องต้น

<!-- grade_id: onehot_test, points: 2, grade: true, solution: false -->
```python
assert len(df_field.columns) == 6
for i in range(6):
    assert df_field.columns[i].startswith('Field_')

```

(PC2.4) จงนำข้อมูลใน `df_field` จากข้อก่อนหน้าซึ่งเก็บผลของการทำ one-hot encoding กับ `EducationField` ไปแทนที่คอลัมน์ `EducationField` เดิมใน `df` แล้วเก็บผลไว้ในตัวแปรใหม่ชื่อ `df_clean` 

นั่นคือ `df_clean` จะต้องมีข้อมูลเหมือน `df` เพียงแค่คอลัมน์ `EducationField` ได้ถูกเปลี่ยนเป็นแบบ one-hot ทั้งนี้ลำดับของคอลัมน์ต่างๆ อาจเปลี่ยนไปจาก `df` ได้และไม่มีผลต่อคะแนน

<!-- grade_id: onehot_replace, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

ตรวจสอบความถูกต้องของ `df_clean` ในเบื้องต้น

<!-- grade_id: onehot_replace_test, points: 3, grade: true, solution: false -->
```python
assert len(df_clean.columns) == 20
```